# Introduksjon til Jupyter Notebook
Programmet du nå er inne i heter *Jupyter Notebook*. Det er et program som lar deg kjøre Python-kode *mens* du skriver den. Det gjør det raskere å skrive korte programmer, og det er også et fantastisk verktøy for å visualisere data som du vil se etter hvert.

## Celler
Jupyter er bygd opp av celler. Celler kan være tekst (som denne cellen) eller kode. Under er det en celle med kode. Man kan se at det er en kode-celle fordi det står 
<div class="input"><div class="prompt input_prompt">
In [ ]:
</div></div>
til venstre for cellen.

For å kjøre en celle med kode setter du markøren inni cellen ved å trykke på den og trykke `ctrl-Enter`. Prøv det på cellen under.

```python
print 'Hello, World!'
```

Til venstre for cellen burde det nå stå
<div class="input"><div class="prompt input_prompt">
In [1]:
</div></div>
`1`-tallet betyr det er den første cellen med kode som er kjørt. Du kan gjøre cellen flere ganger, og da vil tallet øke.

Under cellen burde det nå stå "`Hello, World!`" som er det koden produserte.

<div class="alert alert-danger">Hvis du dobbelttrykker på en tekstcelle vil teksten forandre seg til koden som lager teksten. Dette er ikke viktig for deg. Hvis du ved et uhell dobbelttrykket på en tekstcelle, kan du gjøre den om til vanlig tekst ved å trykke `ctrl-Enter`, akkurat som for celler med kode.</div>

Prøv å skrive "`print 1 + 1`" i cellen under og kjøre den. Den burde skrive ut "`2`".

```python

```

## Labøvelser i Jupyter Notebook
Du skal gjøre deler av labøvelsene i Jupyter Notebook. Da vil det stå instrukser for hva du skal gjøre. I de fleste oppgavene vil du bli hjulpet ved at deler av koden allerede er oppgitt og du bare trenger å fylle ut noen linjer. Under er det et eksempel der man lager en fuksjon som skriver tekst baklengs.

<!-- grade_id: test_reverse, points: 0, grade: false, solution: true -->
```python
def reverse(text):
    # SKRIV DITT SVAR HER
```

Der hvor det står at du skal fylle inn din kode, skal du så klart skrive din besvarelse. Fyll ut cellen over slik at det står
```python
def reverse(text):
    return text[::-1]
```
og kjør koden ved å trykke `ctrl-Enter`. Kjør så cellen under for å undersøke om funksjonen oppfører seg som forventet. Du kan prøve å skrive inn andre ting også.

```python
print reverse('Hello, World!')
```

<div class="alert alert-danger">Hvis du får opp en feilmelding som
<pre><span class="ansi-red-fg">---------------------------------------------------------------------------</span>
<span class="ansi-red-fg">NameError</span>                                 Traceback (most recent call last)
<span class="ansi-green-fg">&lt;ipython-input-X-XXXXXXXXXXXX&gt;</span> in <span class="ansi-cyan-fg">&lt;module&gt;</span><span class="ansi-blue-fg">()</span>
<span class="ansi-green-fg">----&gt; 1</span><span class="ansi-red-fg"> </span><span class="ansi-green-fg">print</span> reverse<span class="ansi-blue-fg">(</span><span class="ansi-blue-fg">'Hello, World!'</span><span class="ansi-blue-fg">)</span>

<span class="ansi-red-fg">NameError</span>: name 'reverse' is not defined

</pre>
så har du glemt å kjøre cellen over.
</div>

## Testfunksjoner
Når du løser programmeringsoppgaver knyttet til labben vil du bli hjulpet ved at det er oppgitt testfunksjoner. Testfunksjonene kjører det du har programmert og sjekker at svaret som kommer ut er fornuftig. Under er det gitt et eksempel på en testfunksjon for `reverse`-funksjonen over.

<div class="alert alert-warning">Hvis testfun
ksjonen ikke får det svaret den forventet, får du en feilmelding som ser slik ut:
<pre><span class="ansi-red-fg">---------------------------------------------------------------------------</span>
<span class="ansi-red-fg">AssertionError</span>                            Traceback (most recent call last)
<span class="ansi-green-fg">&lt;ipython-input-X-XXXXXXXXXXXX&gt;</span> in <span class="ansi-cyan-fg">&lt;module&gt;</span><span class="ansi-blue-fg">()</span>
<span class="ansi-green-fg">----&gt; 1</span><span class="ansi-red-fg"> </span><span class="ansi-green-fg">assert</span> <span class="ansi-cyan-fg">YYY</span> <span class="ansi-blue-fg">==</span> <span class="ansi-cyan-fg">ZZZ</span>

<span class="ansi-red-fg">AssertionError</span>: 

</pre>
</div>

Da må du gå tilbake og prøve å finne feilen i programmet ditt.

<!-- grade_id: palindrome_test, points: 1, grade: true, solution: false -->
```python
assert reverse('abc 123') == '321 cba'
assert reverse('grav ned den varg') == 'grav ned den varg'
```

## Levere
For å levere koden din går du tilbake til *Assigments*-fanen i Jupyter.

Du kan da trykke <button class="btn btn-default btn-xs">Validate</button>-knappen for at Jupyter skal kjøre koden din om igjen.

<div class="alert alert-danger">
Dersom koden din ikke blir validert selv om du kunne ha sverget på at den fungerte er det nok fordi du hadde definert variabler og så slettet definisjonene senere. Variabler Jupyter blir hengende selv om du sletter definisjonene av dem. Du kan trykke <button class="btn btn-default"><i class="fa-repeat fa"></i></button>-knappen øverst som sletter alle variablene. Du må da gjøre all koden igjen som du kan gjøre ved å trykke *Cell - Run all* fra menyen øverst.
</div>

Så kan du levere koden din ved å trykke <button class="btn btn-primary btn-xs">Submit</button>.
