# Lab 1 - Statistikk
I denne øvelsen skal vi lage våre egne funksjoner for å gjøre grunnleggende statistikk på datasett.

## Gjennomsnitt
Det første vi skal se på er gjennomsnitt. Dersom $x_1, x_2, x_3, \ldots, x_n$ er måledata er gjennomsnittet $\bar x$ gitt ved
$$\bar x = \frac1n \sum_{i=1}^n x_i = \frac1n \left(x_1 + x_2 + x_3 + \ldots + x_n\right).$$
Implementer en funksjon `mean(data)` som regner ut og returnerer gjennomsnittet til listen av tall `data` ved å bruke en `for`-løkke som itererer over alle elementene i listen.

<!-- grade_id: mean, points: 0, grade: false, solution: true -->
```python
def mean(data):
    # Her skriver du koden for å regne ut gjennomsnittet.
    # Husk a returnere gjennomsnittet. 
    # SKRIV DITT SVAR HER
    return float(sum(data))/len(data)
```

Kjør testkoden under for å sjekke at koden din er okay.

<div class="alert alert-danger">Hvis alle bortsett fra de to siste testene fungerer burde du undersøke om du har heltallsdivisjon. Husk at i Python 2 vil et heltall delt på et heltall også gi et heltall så <code>1/2=0</code> mens <code>1./2=0.5</code>.</div>

<!-- grade_id: test_mean, points: 1, grade: true, solution: false -->
```python
assert mean([1, 4, 1, 3, 6]) == 3
assert mean((5, -3, 8, -0, 33, -17, 9)) == 5
assert abs(mean([1.5, 3.2, -0.17, 4, 1./3]) - 1.77267) < 1e-5
assert mean([1, 2]) == 1.5
assert abs(mean([1, -1, 1]) - 0.33333) < 1e-5
```

<div class="alert alert-warning">Koden for gjennomsnittet kan skrives enklere uten å bruke en <code>for</code>-løkke. <code>sum( )</code>-funksjonen regner ut summen av elementene i en listen. Klarer du å lage en funksjon for å regne ut gjennomsnittet som er bare én linje?</div>

## Varians og standardavvik
Varians og standardavvik er et mål på spredningen til data. Den empiriske variansen $s^2$ regnes ut ved
$$s^2 = \frac1{n-1} \sum_{i=1}^n \left(x_i - \bar x \right)^2$$
der $\bar x$ er gjennomsnittet regnet ut i forrige oppgave.

Standardavviket $s$ er følgelig $s = \sqrt{s^2}$.

Lag funksjonene `var(data)` og `stdev(data)` som regner ut henholdsvis variansen og standardavviket til listen med tall `data`. Funksjonen `var` *skal bruke funksjonen `mean` fra forrige del og `stdev` skal bruke bruke `var`*.

<!-- grade_id: var, points: 0, grade: false, solution: true -->
```python
def var(data):
    return 1
```

<!-- grade_id: stdev, points: 0, grade: false, solution: true -->
```python
def stdev(data):
    # SKRIV DITT SVAR HER
    raise NotImplementedError
```

<!-- grade_id: var_stdev_test, points: 1, grade: true, solution: false -->
```python
assert abs(var([1, 2, 3, 4]) - 1.66667) < 1e-5

#Check if a function uses another function
def func_uses_func(func1, func2, args, expected, tol=1e-5):
    temp_func1 = func1
    del func1
    try:
        assert abs(func1(args) - expected) < tol
    except NameError:
        pass
    else:
        raise AssertionError(func1.__name__ + ' does not appear to use ' + func2.__name__)
    finally:
        func1 = temp_func1
        del temp_func1
        
func_uses_func(var, mean, [1.6, 3.1, -3.6], 12.36333)
func_uses_func(stdev, var, [1.45, -3.14, 6.7, 9, .620], 4.884391)
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-15-7f80600094d4> in <module>()
----> 1 assert abs(var([1, 2, 3, 4]) - 1.66667) < 1e-5
      2 
      3 #Check if a function uses another function
      4 def func_uses_func(func1, func2, args, expected, tol=1e-5):
      5     temp_func1 = func1

AssertionError:
```

```python

```