# Titreringssimulator

Programmet under plotter en titrerkurve. Det skal regnes ut pH til en løsning når en svak én-protisk syre titreres med sterk base. Plottet skal sammenligne teoretisk utregnet pH og observert pH i forsøk 5. 

Dette er en oppskrift og hjelp til labrapporten du skal levere til lab 5. Det er ikke pålagt at du skal bruke denne hjelpen. Du står fritt til å lage ditt eget program for å gjøre utregningene og lage plottet.


Programmet under er delt inn i flere deler. Du skal:
    - fylle inn observert data fra lab-øvelsen i lister
    - skrive en funksjon som regner ut og returnerer pH-verdi 
    - kalle på funksjonen og lage et array med utregnet pH 
    - plotte både utregnet og observert pH mot tilsatt base i samme plot
    
    
Du bør følge rekkefølgen slik programmet er skrevet for å unngå feilmeldinger. 

Mye av koden er ferdigskrevet. De stedene hvor du skal fylle inn kode er markert slik:
 
```
        ### FORKLARING PÅ HVA SOM SKAL KODES ###
```


Noen steder er det lagt inn ferdig kode for test-funksjoner. Disse hverken kan eller skal endres på, men kjøres når du er ferdig med å fylle inn koden i feltet over. Test-funksjonen vil gi deg tilbakemelding på eventuelle feil i koden du har skrevet. Om test-funksjonen ikke slår ut (altså at den kjører stille) er dette ingen garanti for at du har skrevet alt riktig, men den vil plukke opp og gi feilmelding på noen typiske feil.


Du kjører koden ved å trykke:   shift + enter


<mark><strong>HUSK:</strong></mark> Programmet vil lagre variabler og funksjoner når du kjører bitene av programmet. For å unngå feil bør du nullstille "output"-en ofte. Dette gjør du slik:
    - klikk på "Cell"
    - velg "All Output"
    - klikk på "Clear"
    
    
## Observerte verdier
Fyll inn observert pH fra lab i pH_obs.

Fyll inn volumet tilsatt base ved tidspunkt for avlesning av pH.

<!-- grade_id: Lister, points: 0, grade: false, solution: true -->
```python
pH_obs = []

V_when_obs = []
```

Testfunksjonen under tester om listene dine er like lange. Den vil plukke opp om du har glemt å skrive inn en verdi i den ene listen. Kjør koden når du har skrevet inn alle verdiene i listene over.

<!-- grade_id: test_Lister, points: 0, grade: true, solution: false -->
```python
def test_lists():
    success = len(pH_obs) == len(V_when_obs)
    message = 'The two lists differ in length'
    assert success, message
test_lists()
```

## Funksjon for beregning av pH


Variabelnavnene ```V``` (volum), ```c``` (konsentrasjon) og ```n``` (antall mol) vil gjelde for ```a``` (syren) eller ```b``` (basen). pKa er for den svake syren.


I dette programmet skal du fylle inn likninger for å regne ut pH.
    - første tilfelle: det forekommer ren syre i løsningen. Formel for å regne ut pH for svak syre er brukt.
    - andre tilfelle: mengde base er mindre enn mengde syre. Bufferligningen tas i bruk.
    - tredje tilfelle: mengde base er lik mengde syre. Utregning av pH ved ekvivalenspunkt.
    - fjerde tilfelle: mengde base er større enn mengde syre. Tilsatt base etter ekvivalenspunkt gir pH for sterk base i vann.
    

HINT: numpy har en funksjon som heter ```log10()```.

<!-- grade_id: calc_pH, points: 0, grade: false, solution: true -->
```python
import numpy as np
import matplotlib.pyplot as plt

def calculate_pH(Vb, Va, cb, ca, pKa):
    na = ca*Va        # moles acid
    nb = cb*Vb        # moles base
    Kb = 1e-14/10**(-pKa)
    
    if nb == 0:
        pH = 0.5*(pKa - np.log10(ca))
    elif 0 < nb < na:
        pH = pKa + np.log10( nb/(na - nb) )
    elif abs(na - nb) < 1e-12:
        pH = 14 + 0.5*np.log10( Kb*ca*cb/(ca + cb))
    else:
        pH = 14 + np.log10((nb - na)/(Va + Vb))
    return pH
```

Funksjonen under tester om formlene i ```calculate_pH``` er skrevet riktig. Den vil gi tilbakemelding om hvilken ligning som gir uventet svar. Les derfor eventuell feilmelding. 

Kjør koden når du har skrevet ferdig koden i ```calculate_pH```.

<!-- grade_id: calc_pH_test, points: 4, grade: true, solution: false -->
```python
def test_calculate_pH():
    pKa = 4.75
    ca = 0.15
    cb = 0.15
    Va = 0.1
    
    Vb = [0, 0.05, 0.1, 0.12]
    pH_expected = [2.78695437047, 4.75, 8.81253063169585, 12.1346985739]
    partOf = ['nb = 0', '0 < nb < na', 'nb == na', 'nb > na']
    
    tol = 1e-11
    for i in range(4):
        pH_calc = calculate_pH(Vb[i], Va, cb, ca, pKa)
        success = abs(pH_expected[i] - pH_calc) < tol
        message = 'Unexpected value encountered when %s' %partOf[i]
        assert success, message

test_calculate_pH()
```

## Beregning av pH
Fyll inn dataene for variablene. 

```Vb_array``` inneholder en liste "tilsatt" base. pH for løsning regnes ut for hvert element i arrayet.

<!-- grade_id: lag_array, points: 0, grade: false, solution: true -->
```python
pKa = 4.75
ca = 0.15
cb = 0.15
Va = 0.1

maxVolumeBase = 0.15
N = 100
Vb_array = np.linspace(0, maxVolumeBase, N)

pH_calculated = np.zeros(N)

for i in range(N):
    pH_calculated[i] = calculate_pH(Vb_array[i], Va, cb, ca, pKa)

```

## Plotting av titrerkurve

Koden for plottet er nesten ferdig skrevet. Både den utregnede pH-en og den observerte pH-en vil inkluderes i samme plot. Det er viktig å merke seg at observert pH IKKE er tegnet med en linje, men som punkter. Måten dette er blitt kodet, er ved å sette in 'o' bak punkt-listene til aksene. (```plt.plot(x-akse, y-akse, 'o')```) Du kan selvfølgelig legge til mer informasjon til plottet om du skulle ønske dette. Det eneste du trenger å gjøre er å fylle inn den manglende informasjonen som er markert med ```"---??---"```. Husk å markere teksten som string ved hjelp av anførselstegn.

<!-- grade_id: plot, points: 2, grade: true, solution: true -->
```python
plt.plot(V_when_obs, pH_obs, 'ro', label='observed pH')
plt.plot(Vb_array, pH_calculated, label='calculated pH')
plt.xlabel('volume added base')
plt.ylabel('pH')
plt.legend()
plt.title('Titration')
plt.show()
```

HINT: kan ```V_when_obs``` og funksjonenen ```calculate_pH``` brukes til noen utregninger i labrapporten?

<!-- grade_id: Hint, points: 0, grade: false, solution: true -->
```python
for V_b in V_when_obs:
    print(calculate_pH(V_b, Va, cb, ca, pKa))
```

```python

```