# Ideell versus van der Waals

Dette programmet skal inheholde:
    - en funksjon som regner ut trykk ved å bruke den ideelle gassloven
    - en funksjon som regner ut trykk ved å bruke van der Waals' likning
    - en kodesnutt som leser av van der Waalskonstantene, og printer ut utregnet trykk
 
Mye av koden er allerede skrevet, men steder hvor kode mangler er markert slik:
     # FORKLARING PÅ HVA SOM SKAL IMPLEMENTERES
    
 
 
Det ligger en testfunksjon i dette programmet. Koden til testfunksjonen er ferdig skrevet, og alt du trenger å gjøre er å kjøre koden når du har skrevet ferdig ideal(n, T, V) og van_der_Waals(n, T, V, a, b). Testfunksjonen vil fortelle deg hvilken av funksjonene som eventuelt ikke returnerer forventet verdi. Merk at testfunksjonen ikke alltid vil finne alle feil, og er ingen absolutt garanti på at funskjonen er implementert helt riktig.

Du kjører koden ved å trykke:   shift + enter


HUSK: Programmet vil lagre variabler og funksjoner når du kjører bitene av programmet. For å unngå feil bør du nullstille "output"-en ofte. Dette gjør du slik:
    - klikk på "Cell"
    - velg "All Output"
    - klikk på "Clear"


Gasskonstanten er oppgitt nedenfor. Denne skal ikke redigeres, men du må kjøre koden for at programmet aktiverer og dermed husker variabelen.

```python
R = 0.0821
```

I neste celle skal du skrive en funksjon som regner ut trykk ved å bruke den ideelle gassloven: $$ P = \frac{nRT}{V} $$ Hvor P er trykk, T er temperatur målt i Kelvin, n er antall mol og R er gasskonstanten.Gasskonstanten er definert ovenfor som global variabel, og vil derfor kunne brukes så lenge du har kjørt cellen over en gang etter du nullstilte output. Funksjonen tar inn n, T og V som parameter.

<!-- grade_id: Ideell, points: 0, grade: false, solution: true -->
```python
def ideal(n, T, V):
    P =    #fyll in formel her
    return P
```

I neste celle skal du skrive en ny funksjon som regner ut trykk ved å bruke van der Waals' likning: $$ P = \frac{nRT}{V - nb} - \frac{an^2}{V^2}$$ Hvor a er konstanten for korreksjon av trykk og b er konstanten for korreksjon av volum. 

I tillegg til n, T og V, tar funksjonen også inn a og b som parameter.

<!-- grade_id: van_der_Waals, points: 0, grade: false, solution: true -->
```python
def van_der_Waals(n, T, V, a, b):
    P =    #fyll inn formel her
    return P
```

Neste celle er en ferdigskrevet testfunksjon. Kjør denne når du er ferdig med å implementere ideal(n, T, V) og van_der_Waals(n, T, V, a, b).

<!-- grade_id: tester, points: 4, grade: true, solution: false -->
```python
n = 1
V = 1.30
T = 300
tol = 1e-8

def test_ideal():
    expected = 18.9461538462
    computed = ideal(n, T, V)
    success = abs(expected - computed) < tol
    message = 'Funksjonen ideal gir uventet svar!'
    assert success, message
    
def test_van_der_Waals():
    expected = 19.2778523533
    computed = van_der_Waals(n, T, V, 0.034, 0.0237)
    success = abs(expected - computed) < tol
    message = 'Funksjonen van_der_Waals gir uventet svar!'
    assert success, message

test_ideal()

test_van_der_Waals()
```

I neste celle skal du oppgi antall mol, volum og trykk. Du kan selv velge passende verdier.

Filen som skal leses er på denne formen:
```
Gas     a(atm L^2/mol^2)    b(L/mol)
He           0.034           0.0237
Ne           0.211           0.0171
.              .                .
.              .                .
.              .                .
NH3          4.17            0.0371
(values from table 5.3, Chemical Principles)
```
Fyll inn slik som kommentert nedenfor. Være obs på at verdier hentet ut fra fil alltid vil være string. Du må derfor caste om a og b til float. Det kan gjøres slik:
```
        a = float(avlest_verdi_for_a)
```

Print-statments skal inneholde navnet på gassen, og hva trykket ble. Trykket skal formatteres slik at du begrener antall siffer som printes ut. Koden for første print-statment er allerede skrevet slik at du kan se på denne om du er usikker på hvordan dette gjøres.

Programmet vil åpne filen og lese første linje uten å lagre informasjonen. Resten av linjene hentes ut gjennom en for-loop. Denne avsluttes når if-testen kjenner igjen siste linje, som begynner med parentes. Fyll in nødvenig kode som beskrevet i for-loopen.


<!-- grade_id: skriv_ut_trykk, points: 4, grade: true, solution: true -->
```python
n =    #oppgi antall mol
V =    #oppgi volum
T =    #oppgi temperatur

ideell = ideal(n, T, V)
print 'Ideell : %.1f atm' %(ideell, 'ideell')

with open('vanDerWaalsConstants.txt', 'r') as infile:
    infile.readline()
    for line in infile:
        if line[0] == '(':
            break
        newLine =    #bruk split() for å dele opp line til en liste
        gas =        #hent nullte element fra newLine
        a =          #hent første element fra newLine og cast til float
        b =          #hent andre element fra newLine og cast til float
        real =       #kall på van_der_Waals, pass på at du setter inn parameter i riktig rekkefølge
        #print ut som vist for ideell

```