# Reelle gasser

Programmet du skal skrive i denne oppgaven, skal regne ut trykk ved å bruke van der Waals ligning. For å kunne sammenligne avviket med en ideell gass, skal du holde alle andre parametere konstant i utregningene.


Dette programmet skal inheholde:
    - en funksjon som regner ut trykk ved å bruke van der Waals' likning
    - en kodesnutt som leser av van der Waalskonstantene fra filen vanDerWaalsConstants.txt, og printer ut utregnet trykk for hver gass som er oppgitt i filen
 
Mye av koden er allerede skrevet, men steder hvor kode mangler er markert slik:
     ### FORKLARING PÅ HVA SOM SKAL IMPLEMENTERES ###
    
 
 
Det ligger en testfunksjon i dette programmet. Koden til testfunksjonen er ferdig skrevet, og alt du trenger å gjøre er å kjøre koden når du har skrevet ferdig `van_der_Waals(n, T, V, a, b)`. Testfunksjonen vil fortelle deg om funksjonen eventuelt ikke returnerer forventet verdi. Merk at testfunksjonen ikke alltid vil finne alle feil, og er ingen absolutt garanti på at funksjonen er implementert helt riktig.

Du kjører koden ved å trykke:   shift + enter


**HUSK:** Programmet vil lagre variabler og funksjoner når du kjører bitene av programmet. For å unngå feil bør du nullstille "output"-en ofte. Dette gjør du slik:
    - klikk på "Cell"
    - velg "All Output"
    - klikk på "Clear"

I neste celle skal du skrive en ny funksjon som regner ut trykk ved å bruke van der Waals' likning: 

$$ P = \frac{nRT}{V - nb} - \frac{an^2}{V^2}$$ 

Hvor $a$ er konstanten for korreksjon av trykk og $b$ er konstanten for korreksjon av volum.
```

```

`van_der_Waals` tar inn `n`, ` T`, `V`, `a` og `b` som parameter. Fyll inn utrykket for $P$.

<!-- grade_id: van_der_Waals, points: 0, grade: false, solution: true -->
```python
R =  0.0821

def van_der_Waals(n, T, V, a, b):
    P = n*R*T/(V - n*b) - a*(n/V)**2
    return P
```

Neste celle er en ferdigskrevet testfunksjon som tester de funkjsonene du har skrevet ovenfor. Kjør denne når du er ferdig med å implementere van_der_Waals(n, T, V, a, b).

<!-- grade_id: tester, points: 0, grade: true, solution: false -->
```python
def test_van_der_Waals():
    n = 1
    V = 1.30
    T = 300
    tol = 1e-8
    expected = 19.2778523533
    computed = van_der_Waals(n, T, V, 0.034, 0.0237)
    success = abs(expected - computed) < tol
    message = 'Funksjonen van_der_Waals gir uventet svar!'
    assert success, message

test_van_der_Waals()
```

Vis at van der Waals ligning er lik den ideelle gassloven hvis man setter $a$ og $b$ til å være null. 
```

```
I neste celle skal du oppgi antall mol, volum og trykk. Du kan selv velge passende verdier. Kjør så koden for å lagre variablene.

<!-- grade_id: ideell, points: 0, grade: true, solution: true -->
```python
moles = 1
volume = 1.3
temperature = 300

gas = 'Ideal'
ideal = van_der_Waals(moles, temperature, volume, 0, 0)
```

Filen som skal leses er på denne formen:
```
Gas     a(atm L^2/mol^2)    b(L/mol)
He           0.034           0.0237
Ne           0.211           0.0171
.              .                .
.              .                .
.              .                .
NH3          4.17            0.0371
(values from table 5.3, Chemical Principles)
```
Fyll inn slik som kommentert nedenfor. Være obs på at verdier hentet ut fra fil alltid vil være av type string. Du må derfor konvertere om $a$ og $b$ til float. Det kan gjøres slik:
```
        a = float(avlest_verdi_for_a)
```

Print-statments skal inneholde navnet på gassen, og hva trykket ble. Trykket skal formatteres slik at du begrenser antall siffer som printes ut. Koden for første print-statment er allerede skrevet slik at du kan se på denne om du er usikker på hvordan dette gjøres.

Programmet vil åpne filen og lese første linje uten å lagre informasjonen. Resten av linjene hentes ut gjennom en for-løkke. Denne avsluttes når if-testen gjenkjenner siste linje, som begynner med parentes. Fyll in nødvenig kode som beskrevet i for-løkken.



<!-- grade_id: skriv_ut_trykk, points: 0, grade: true, solution: true -->
```python
print('%s : %.1f atm' %(gas, ideal))

with open('vanDerWaalsConstants.txt', 'r') as infile:
    # getting rid of first line\n",
    infile.readline()
    for line in infile:
        # recognizing the last line in the file\n",
        if line[0] == '(':
            break
        newLine = line.split()
        gas = newLine[0]
        a = float(newLine[1])
        b = float(newLine[2])
        real = van_der_Waals(moles, temperature, volume, a, b)
        print('%s : %.1f atm' %(gas, real))
```